Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ahmed"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 5.1

## Question 1.

### Question 1a.

Write the worst sort function in the world, `worst_sort`. This function takes a list, and then randomly shuffles it until it happens to be in sorted order. Once it is in sorted order then the list is returned.


<!-- grade_id: cell-ead9c74054a1c5da, points: 0, grade: false, solution: true -->
```python
import numpy as np
def worst_sort(A):
    """
    Sort A in ascending order by randomly shuffling its 
    elements until they are in order.
    
    Input:
    - A: list of numerical values
    
    Output:
    - A: Sorted list
    """
    # YOUR CODE HERE
    np.random.shuffle(A)
    for i in range(len(A)-1):
        if A[i]> A[i+1]:
            worst_sort(A)
    return(A)
    raise NotImplementedError()
```

```python
test=[5,3,4,6,10,48]
worst_sort(test)
```

Output:
```
[3, 4, 5, 6, 10, 48]
```

<!-- grade_id: cell-ff4db3a3e8b04515, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 1b.
What is the best case complexity of this algorithm?

<!-- grade_id: cell-afac80171b3b6232, points: 0, grade: true, solution: true -->
The best case complexity when the algorithm sorts the list from the first time using the shuffle function. In this case, no recursion will happen and the complexity will be theta(n).

### Question 1c.

What is the average case complexity?


<!-- grade_id: cell-9d52d34daa1e3478, points: 0, grade: true, solution: true -->
The worst case complexity is that the algorithm will run forever, assuming that random.shuffle is completly random. The average case is that the algorithm will run for certian time then it will reach to the sorted list. I think the average case complexity will be O(n^k), where k is the number of the shuffles.

### Question 1d.

For what size lists is this a feasible method?

<!-- grade_id: cell-ab40f08768579798, points: 0, grade: true, solution: true -->
The worst_sort algorithm is a feasible method for small lists as there is a higher probablity that shuffle function will generate a sorted list. I think a feasible size can be a list of 2 items where the list has 50% to be sorted in the right way. 

### [Optional] Question 1e.

Can you think of an even worse sorting method? In such a case, what would its complexity be? How big a list could you sort?

<!-- grade_id: cell-ac05806942341926, points: 0, grade: true, solution: true -->
A worse sorting method will be adding an additional sorting function like bubble sorting to compare the results from the random.suffle function. 
For example:
    np.random.shuffle(A)
    while A != bubblesort(A):
       worst_sort(A)
    return(A)

In this case, the complexity can be O(n^2k) as the complexity for the bubble sort is already O(n^2). 

## Question 2.
Approximate median finder. Given a list and δ (a number between 0 and 0.5), the approximate median finder function returns a number that is guaranteed to lie between the (50-δ/2)% and (50+δ/2)% percentiles. Implement such a function by randomly selecting an element from the list and testing whether or not it lies within the bounds. If it doesn’t then retry with a new random element, but only a limited number of retries are allowed (you decide the maximum number of retries.)

### Question 2a.
Complete the following function


```python
def selectionsort(A):
    for i in range(len(A)-1):  
        minidx = i 
        for j in range(i+1, len(A)): 
            if A[minidx] > A[j]: 
                minidx = j         
        A[i], A[minidx] = A[minidx], A[i]
    return A
```

<!-- grade_id: cell-2ab65512a6148d3c, points: 0, grade: false, solution: true -->
```python
import math
def check_approx_med(A, num, delta):
    """
    Given a list and a number in a list, determine whether it is within (50-delta/2)% and 
    (50+delta/2)% percentiles.
    X% percentile of a list is defined to be the smallest number in the list that is as large as at least
    X% numbers in the list. 
    
    Input:
    - A: list of numerical values
    - num: the number of interest
    - delta: the error, lies between 0 and 0.5
    
    Output:
    - isApproxMed: a boolean value indicating whether num is within the bound
    """
    # YOUR CODE HERE
    lbound = (50-delta/2)/100
    Ubound= (50+delta/2)/100
    AS=selectionsort(A)
    lpercentile= math.ceil((len(A))*lbound)
    Upercentile= math.ceil((len(A))*Ubound)
    
    if len(A)==1:
        return False
    elif (num>=AS[lpercentile-1] and num<=AS[Upercentile-1]):
        return (True)
    else:
        return False
    raise NotImplementedError()
```

```python
test1=[0,1]
test2=[0]
[5,2,1,3,4]
print (check_approx_med(test1,0,.25))
print (check_approx_med(test2,0,.25))
print (check_approx_med([5,2,1,3,4],3,.25))
```

Output:
```
True
False
True
```

<!-- grade_id: cell-df56d7433abb6517, points: 1, grade: true, solution: false -->
```python
assert(check_approx_med([0,1], 0, .25) == True)
assert(check_approx_med([0], 0, .25) == False)
```

### Question 2b.
Complete the following function that makes use of `check_approx_med` above.


<!-- grade_id: cell-e19685b545147dde, points: 0, grade: false, solution: true -->
```python
import random
def approx_med_finder(A, delta):
    """
    Given a list, find a number in the list that is between 50-delta/2% and 50+delta/2% percentiles
    
    Input:
    - A: list of numerical values
    - delta: the error, lies between 0 and 0.5
    
    Output:
    - num: the approximated median, if it is found within 100//delta trials (this is one possible 
    maximum number of retries. While you are encouraged to play around with another limits, the 
    submitted work must use this number.)
    - None: finding failed, if nothing found within 100//delta trials

    Note: the 100//delta is chosen here as it is the expected number of trials for the first successful finding to occur (using geometric distribution)
    """
    # YOUR CODE HERE
    for i in range(int(100//delta)): 
        num=random.choice(A)
        if check_approx_med(A,num,delta):
            return num
    return print ("finding failed, if nothing found within 100//delta trials")
    raise NotImplementedError()
```

```python
test= [5,2,1,3,4]
test2= [6,2,4,5,6,8,9,10]
approx_med_finder(test,0.25)
approx_med_finder(test2,0.25)
```

Output:
```
6
```

### Question 2c.

What is the probability of failure in each random trial? What is the probability of failure after all the allowed random trials? Does this scale with δ or N (the number of elements in a list)?


<!-- grade_id: cell-efe7ab76f6a2546a, points: 0, grade: true, solution: true -->
The probability of success in each random trail is 1/n as we can have one median in each set. The probability of failur in each random trail will be (n-1)/n. After finishing all the allowed random trials, this mean that we have finished all the expected number of trials for the first successful finding to occur. Thus, the probability of failure will be approximatly near to 100%. The probability of failure in each random trial scale with N the number of elements in the list. However, the probability of failure after all the allowed random trials is indpendent to the δ value as it's the expected number of trails regardless the delta value.

### Question 2d.
Analyze the expected runtime of `approx_med_finder`. Note that because the function uses `check_approx_med`, you will most likely need to analyze the runtime of that function, too.

<!-- grade_id: cell-86b199798477fc2a, points: 0, grade: true, solution: true -->
The approx_med_finder has a for loop with range 100//delta, which means it will run 100//delta in its worst case of not finding the median. check_approx_med is a function of n as its complexity depends on n (number of items in the list). With the selection sort function in it, check_approx_med will have a complexity of O(n^2) as it will dominate the complexity. Thus, for both functions the approx_med_finder will be O(n^2) in its worst case. 


```python

```