Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ahmed"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 3.2

## Question 1.
Given the array `H=[39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]`, perform the following operations:
1. Draw the corresponding binary tree of H. Is the binary tree a valid max heap? Explain your answer.
2. Using as a model the drawing examples illustrated in Figure 6.2 of Cormen et al.,  draw a step-by-step transformation of the array above into a valid max heap. 
3. Now that you have obtained a valid max heap, write out the corresponding array that stores the valid max-heap.

Use as many cells as you wish for this question.

<!-- grade_id: cell-b1f821b034b44619, points: 0, grade: true, solution: true -->
![Question1](3_2_1.jpg)

![Question2](3_2_2.jpg)

 The corresponding array with a valid max heap will be H=[92,85,85,76,39,30,49,16,49,15,21,7,29,31,28]

## Question 2. 
Consider the following questions on the $MAX-HEAPIFY$ operation.
### Question 2a.

In the pseudocode of $MAX-HEAPIFY$ (Cormen et al., p.154, or you can view it [here](https://drive.google.com/open?id=1e_3jsX4-qQCfZXKMok_T6LPFh9FwtmT5)), what does A.heap-size mean and what is the idea behind the local variable largest? 


<!-- grade_id: cell-06106b81a909b9e6, points: 0, grade: true, solution: true -->
The A.heap-size, in this case, is the number of the items in the list and it makes sure that the l and r are still in the range of the list.

The idea behind the local variable largest is saving the larger number of the three nodes. (The parent, the right and the left). After saving the largest number, the algorithm re-assign it to the parent to make sure that the list will achieve the max-heap property. 

### Question 2b.
The functions $LEFT(i)$ and $RIGHT(i)$, lines 1 and 2 in the $MAX-HEAPIFY$ pseudocode, return the array index of the left and right child, respectively, of a node in a binary tree. From reading Section 6.1, you know that the input to both functions is an integer number, $i$, which corresponds to the array index of the parent node in the array. Review Section 6.1 for more information. Write a Python implementation of the functions $LEFT(i)$ and $RIGHT(i)$ by filling in the cells below.

<!-- grade_id: cell-2efd4017321c8d48, points: 0, grade: false, solution: true -->
```python
def left(i):
    # YOUR CODE HERE
    return 2*i+1     #The left index is different in python as pythons starts from 0, thus -> 2i+1 
    raise NotImplementedError()
print(left(1))
```

Output:
```
3
```

<!-- grade_id: cell-507ddbab62e6a32c, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

<!-- grade_id: cell-b7e5ceeacaeaf1b8, points: 0, grade: false, solution: true -->
```python
def right(i):
    # YOUR CODE HERE
    return 2*i+2   #The right index is different in python as pythons starts from 0, thus -> 2i+2 
    raise NotImplementedError()
print(right(1))
```

Output:
```
4
```

<!-- grade_id: cell-c24867515653788a, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 2c.
Write a Python implementation of the MAX-HEAPIFY operation using the pseudocode above, and your newly written functions, `left` and `right`.

<!-- grade_id: cell-52289253a243341b, points: 0, grade: false, solution: true -->
```python
A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
def heapify(heap, i):
    """
    Inputs:
    - heap: a list of floats. Assume that the heap size is the length of the heap.
    
    No output is needed. This function should modify (if necessary) heap in-place.
    """
    l = left(i)
    r = right(i)
    temp=0
    if (l < len(heap) and heap[l]>heap[i]): #If statment to make sure 1- that the index is in the Heapsize, 2- to compare which node is bigger the left or the parent
        largest=l
    else:             #line 13m, 15 and 17 are assigning the largest node to the storage variable "largest"
        largest=i
    if  (r < len(heap) and heap[r]>heap[largest]): #If statment to make sure 1- that the index is in the Heapsize, 2- to compare which node is bigger the right or the parent
        largest = r
    if largest != i:
        temp=heap[largest]           #if the right or the left nodes is bigger than the parent, exchange the values of the nodes
        heap[largest]=heap[i]
        heap[i]=temp
        heapify(heap, largest)       #recurrence line to make sure that the same process will happen again but a level lower this time. The largest index (after the exchange) is now the parent  
    raise NotImplementedError()
```

<!-- grade_id: cell-7f63745ff34b881b, points: 1, grade: true, solution: false -->
```python
A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
heapify(A,0)
assert(A == [85, 49, 85, 16, 39, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28])

```

```python

```

## Question 3. 
Next, write a Python implementation of the BUILD_MAX_HEAP operation using the pseudocode provided in Section 6.3 of Cormen et. al. Test your Python implementation using the array in problem 1, and make sure your Python codes produce a valid max heap.

<!-- grade_id: cell-9237da6268a5e3eb, points: 0, grade: false, solution: true -->
```python
def build_max_heap(A):
    """
    Input:
    - A: a list of floats.
    
    No output is needed. The function should turn A into a valid max heap, in-place.
    """
    # YOUR CODE HERE
    root=(len(A)//2)         #Choosing all the roots/parents of the tree
    for i in range(root,-1,-1):  #runing the heapify on all the parents of the tree
        heapify(A,i)
    raise NotImplementedError()

```

<!-- grade_id: cell-5dc147b837da3f7e, points: 1, grade: true, solution: false -->
```python
A = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]
build_max_heap(A)
assert(A == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1])
```

```python
#In this cell, I tested if the algorithm/code is working. And, I successfully achieved the right answers 

A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
B = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]

def build_max_heap(A):
    """
    Input:
    - A: a list of floats.
    
    No output is needed. The function should turn A into a valid max heap, in-place.
    """
    # YOUR CODE HERE
    root=(len(A)//2)
    for i in range(root,-1,-1):
        heapify(A,i)
    return(A)
print (build_max_heap(A))
print (build_max_heap(B))

```

Output:
```
[92, 85, 85, 76, 49, 30, 49, 16, 39, 15, 21, 7, 29, 31, 28]
[16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
```

## Question 4. 

Lastly, write Python implementations of the $MIN-HEAPIFY$ and $BUILD-MIN-HEAP$ operations for a min heap data structure. You can use your $MAX-HEAPIFY$ and $BUILD-MAX-HEAP$ Python function as models, just remember that the latter two functions support operations for the max heap data structure. Test your Python implementation of $BUILD-MIN-HEAP$ using the array in problem 1, and make sure your Python codes produce a valid min heap. 

<!-- grade_id: cell-d86943f127876fc8, points: 0, grade: false, solution: true -->
```python
def min_heapify(heap, i):
    """
    Inputs:
    - heap: a list of floats. Assume that the heap size is the length of the heap.
    
    No output is needed. This function should modify (if necessary) heap in-place.
    """
    # YOUR CODE HERE

    l = left(i)
    r = right(i)
    temp=0
    if (l < len(heap) and heap[l]<heap[i]):  #If statment to make sure 1- that the index is in the Heapsize, 2- to compare which node is smaller the left or the parent
        smallest=l
    else:                     #line 14m, 16 and 18 are assigning the largest node to the storage variable "smalles"
        smallest=i
    if  (r < len(heap) and heap[r]<heap[smallest]): #If statment to make sure 1- that the index is in the Heapsize, 2- to compare which node is smaller the right or the parent
        smallest = r
    if smallest != i:
        temp=heap[smallest]             #if the right or the left nodes is smaller than the parent, exchange the values of the nodes
        heap[smallest]=heap[i]
        heap[i]=temp
        min_heapify(heap, smallest)     #recurrence line to make sure that the same process will happen again but a level lower this time. The "smalles" index (after the exchange) is now the parent
    raise NotImplementedError()
```

<!-- grade_id: cell-6efa45e3956ad0b8, points: 0, grade: false, solution: true -->
```python
def build_min_heap(A):
    """
    Input:
    - A: a list of floats.
    
    No output is needed. The function should turn A into a valid min heap, in-place.
    """
    root=(len(A)//2)   #Choosing all the roots/parents of the tree   
    for i in range(root,-1,-1):   #runing min_heapify on all the parents of the tree
        min_heapify(A,i)
    raise NotImplementedError()
```

<!-- grade_id: cell-7e64941402f819c2, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

```python
#In this cell, I tested if the algorithm/code is working. And, I successfully achieved the right answers 
A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
B = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]

def  build_min_heap(A):
    """
    Input:
    - A: a list of floats.
    
    No output is needed. The function should turn A into a valid max heap, in-place.
    """
    # YOUR CODE HERE
    root=(len(A)//2)
    for i in range(root,-1,-1):
        min_heapify(A,i)
    return(A)
print (build_min_heap(A))
print (build_min_heap(B))
```

Output:
```
[7, 15, 28, 16, 21, 29, 31, 92, 76, 49, 85, 30, 85, 39, 49]
[1, 2, 3, 4, 7, 9, 10, 14, 8, 16]
```

```python

```