Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ahmed"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 2.2

## Question 1 (Exercise 3.1-3 of Cormen, et al. )
Explain why the statement, "The running time of algorithm A is at least $O(n^2)$," is meaningless.


<!-- grade_id: cell-437d61c1420d4f99, points: 0, grade: true, solution: true -->
We use O-notation to give an upper bound on a function, to within a constant factor. Thus, it's meaningless to use al least with the O-notation as it describe the upper bound of that algorithm (asymptotic upper bound).  

## Question 2 (Exercise 3.1-4 of Cormen, et al. )

Is $2^{n+1}=O(2^n)$? Is $2^{2n}=O(2^n)$?

<!-- grade_id: cell-6a0c4ee19dadfddf, points: 0, grade: true, solution: true -->
The first part is possible as there is always f(n) =< cg(n) 
The second part is impossible as sometimes we can't achieve this condition where  f(n) =< cg(n) 

## Question 3.
Write a function in Python that solves the maximum-subarray problem using a brute-force approach. Your Python function must:
* Take as Input an array/list  of numbers
* Produce the following Output: 
    * the start and end indices of the subarray containing the maximum sum.
    * value of the maximum subarray (float)


<!-- grade_id: cell-527e6532d6992fd0, points: 0, grade: false, solution: true -->
```python
def bruteforce_max_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    maxsum = -10000000000000  #setting a negative infinity number as the maximum sum to even consider the negative values in the market 
    sindex =0                 #initiating the start and the end indexes 
    eindex =0
    
    
    #building two nexted loops to consider all the subarrays in the list 
    
    for i in range(len(A)):
        summ= 0
        for j in range(len(A)):
            if j >= i:
                summ = summ + A[j]
                if summ > maxsum:     #Saving only the maximum subarray from the two nested loops + saving the start and the end indexes of this list
                    maxsum = summ
                    sindex= i
                    eindex= j
            
    return ((sindex,eindex,maxsum))  
    raise NotImplementedError()

    #testing the code and achieving the right expected values 
print(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
print(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))    

```

Output:
```
True
True
```

<!-- grade_id: cell-a28a56466c9537ff, points: 1, grade: true, solution: false -->
```python
assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))
```

```python

```

## Question 4. 
Test your Python maximum-subarray function using the following input list (from Figure 4.3 of Cormen et al.):  
`A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7] `

If your Python implementation is correct, your code must return: 
* 43 - which is the answer to the maximum subarray problem, and 
* <7, 10> -the start and the end indices of the max subarray. 



<!-- grade_id: cell-e4a632ce0edc1697, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]
print (bruteforce_max_subarray(A)) #the code worked perfectly :) 
raise NotImplementedError()
```

Output:
```
(7, 10, 43)
NotImplementedError                       Traceback (most recent call last)
<ipython-input-10-bb7eed715753> in <module>
      2 A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]
      3 print (bruteforce_max_subarray(A))
----> 4 raise NotImplementedError()

NotImplementedError:
```

## Question 5. Asymptotic notation. 
Complete the following table using the asymptotic notation that best describes the problem. For example, if both $O(n^3)$ and $O(n)$ are possible for an algorithm, the answer is $O(n)$ because the function $f(n) = O(n)$ provides a tighter and more accurate fit; if both $O(n)$ and $\Theta(n)$ are possible, the correct answer is $\Theta(n)$ because $\Theta(n)$ provides both information about the upper and lower bound, thus it is more accurate than $O(n)$.

You should copy the following table and paste and edit it in the cell below. 

Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort |  |
Selection sort |  |
Bubble sort |  | 
Finding maximum subarray |  |

<!-- grade_id: cell-247c51df276b622e, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort | $O$(n^2)  | $\Theta$(n^2)
Selection sort | $O$(n^2) | $\Theta$(n^2)
Bubble sort | $O$(n^2) | $\Theta$(n^2)
Finding maximum subarray (Brute Force) | $O$(n^2)  |  $\Theta$(n^2)
Merge sort | $O$ nlog(n) | $\Theta$ nlog(n)


**Finding maximum subarray (Brute Force) can be $O$(n^3) depends on how many nested loops do we have in the code**

## [Optional] Question 6. 
How can you change this code to make it find the minimum-subarray?

<!-- grade_id: cell-f4553f71858d1bc4, points: 0, grade: false, solution: true -->
```python
def bruteforce_min_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    minsum = 10000000000000  #setting a positive infinity number as the maximum sum to even consider the negative values in the market 
    sindex =0                 #initiating the start and the end indexes 
    eindex =0
    
    
    #building two nexted loops to consider all the subarrays in the list 
    
    for i in range(len(A)):
        summ= 0
        for j in range(len(A)):
            if j >= i:
                summ = summ + A[j]
                if summ < minsum:     #Saving only the minimum subarray from the two nested loops + saving the start and the end indexes of this list
                    minsum = summ
                    sindex= i
                    eindex= j
            
    return ((sindex,eindex,minsum))  
    raise NotImplementedError()

    #testing the code and achieving the right expected values 
print(bruteforce_min_subarray([-2,1,-1,2,-5]))
print(bruteforce_min_subarray([-2, -5, 6, -2, -3, 1, 5, -6]))  
print(bruteforce_min_subarray([1]*10)[0] == bruteforce_min_subarray([1]*10)[1])
print(bruteforce_min_subarray([1]*10)[2] == 1)

# The code worked just right :))) 

```

Output:
```
(0, 4, -5)
(0, 1, -7)
True
True
```

<!-- grade_id: cell-4263c7494a5f09d3, points: 0, grade: true, solution: false -->
```python
assert(bruteforce_min_subarray([1]*10)[0] == bruteforce_min_subarray([1]*10)[1])
assert(bruteforce_min_subarray([1]*10)[2] == 1)
```