Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ahmed"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 3.1

## Question 1. 
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
A = [-2, -3, 4, -1, -2, 1]
def bruteforce_max_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    maxsum = -float("inf")  #setting a negative infinity number as the maximum sum to even consider the negative values in the market 
    sindex =0                 #initiating the start and the end indexes 
    eindex =0
    
    
    #building two nexted loops to consider all the subarrays in the list 
    
    for i in range(len(A)):
        summ= 0
        for j in range(len(A)):
            if j >= i:
                summ = summ + A[j]
                if summ > maxsum:     #Saving only the maximum subarray from the two nested loops + saving the start and the end indexes of this list
                    maxsum = summ
                    sindex= i
                    eindex= j
            
    return ((sindex,eindex,maxsum))  
    raise NotImplementedError()
print (bruteforce_max_subarray(A))

```

Output:
```
(2, 2, 4)
```

## Question 2. 
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE
Based on the algorithm, we already saved the maximum subarray and the indexes for the start and the end of it. So, I think there is a way to compare these numbers with the new added extra number in the array instead of re-running the entire code. 

## Question 3. 

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
The statement is FALSE
The algorithm works based on comparing the maximum sum of each subarray in the list. If the maximum subarry in B was different than the maximum subarry in A, then there is must be a number that was added in the list to change that. In other words, the newly appended element must be the reason to change the find a new maximum subarray. 

## Question 4.
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
def incremental_max_subarray(x, mx):
    """ 
    Inputs:
    - x: a NON-EMPTY list of numbers (e.g., the array B in the first two questions above). 
        * If x has 1 element: returns the value of the element regardless of the value of mx
    - mx: the maximum subarray of x excluding its last element (i.e., compute the 
    maximum subarray of the input array x considering only its first len(x) - 1 elements)
    
    Output: The maximum subarray of the array x.
    
    For example, using the array B in question 2, the result of incremental_max_subarray(B, 4) 
    is 10 (10 = 8 + 1 - 2 -1 + 4). 
    """
    # YOUR CODE HERE
    nsum = 0  #Creating a storage for the sum of each subarray starting with the last item in the list
    for i in range( len(x) - 1, -1, -1):   #iterating backwards and putting the range from the last index to 0. *Put -1 as the last element in range is not included
        nsum = nsum + x[i]                 #Calculating the sum for each subarray that does include 8
        if nsum > mx:                      #Comparing all the new sum of the new subarray to the maximum subarray and saving the ones with the highest sum
            mx = nsum
    return (mx)                            #returing the maximum sum of the whole array
    raise NotImplementedError()
print(incremental_max_subarray(B, 4) == 10)
print(incremental_max_subarray(B, 4))
```

Output:
```
True
10
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

<!-- grade_id: cell-bacebe71191cbd0f, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
C = [-2.0, -3.0, 4.7, -1.0, -2.5, 1.0, 8.5]
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A to produce the value of the maximum
    subarray of A.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: flaot, the sum of the maximum subarray of A
    """
    # YOUR CODE HERE
    maxsum= -float("inf")   #Setting the sum of the maximum subarray to be -inf so we can all sums that are bigger than this
    length= len(A)
    for i in range(length): #iterating through the orginal loop 
        incremental_max_subarray(A, maxsum) #With each iteration we find the maximum subarray that includes the last number in the list
        maxsum=incremental_max_subarray(A, maxsum) #saving the maximum sum and then we will later compare it with other sums inside the incremental function to find the global maximum subarray
        A=A[0:-1]                           #we change the list and ignore the last number in it because we already found the maximum sum of any subarray that has this number
    return(maxsum)    
    raise NotImplementedError()
print(max_subarray(A))
print(max_subarray(B))
print(max_subarray(C))
```

Output:
```
4
10
10.7
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

```python
for i in range(len(A)):
    print(A[0:-i])
```

Output:
```
[]
[-2, -3, 4, -1, -2]
[-2, -3, 4, -1]
[-2, -3, 4]
[-2, -3]
[-2]
```

### Part 2.
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
The solution for the recurrence tree in the divide and conquer method is T(n)= theta(nlog(n)).
From my understanding of this method, the T(n) = theta (n^2) as it has two loops inside of it. 
Thus, I will say that the divide-and-conquer approach yields an algorithm that is asymptotically faster than the incremental_max_subarray method. (Logarithmic functions are better than quadratic functions)